```scala
import $file.hw2stdlib_new
import hw2stdlib_new._
```

Output:
```
<console>:90: error: not found: value hw2stdlib_new
       import hw2stdlib_new._
              ^
```

```scala

// Nats

sealed trait Nat
case object Zero extends Nat
case class Succ(pred : Nat) extends Nat

val one = Succ(Zero)
val two = Succ(one)
val three = Succ(two)
val four = Succ(three)
val five = Succ(four)
val six = Succ(five)
val seven = Succ(six)
val eight = Succ(seven)
val nine = Succ(eight)
val ten = Succ(nine)

def plus(x : Nat, y : Nat) : Nat = x match {
    case Zero    => y
    case Succ(x) => Succ(plus(x, y))
}

def nat_to_int(x : Nat) : Int = x match {
    case Zero => 0
    case Succ( x ) => 1 + nat_to_int(x)
}

def print_nat(x : Nat) : String = nat_to_int(x).toString

def mult(x : Nat, y : Nat) : Nat = x match {
    case Zero    => Zero
    case Succ(x) => plus(mult(x,y), y)
}


// Booleans

sealed trait Bool
case object True extends Bool
case object False extends Bool

def t = True
def f = False

def id(x : Bool) : Bool = x

def not(x : Bool) : Bool = x match {
    case True => False
    case False => True
}

def and(x : Bool, y : Bool) : Bool = (x,y) match {
    case (True, True) => True
    case _ => False
}

def or(x : Bool, y : Bool) : Bool = (x, y) match {
    case (False, False) => False
    case _              => True
}

def xor(x : Bool, y : Bool) : Bool = (x, y) match{
    case (True, False) => True
    case (False, True) => True
    case _ => False
}

def nand(x : Bool, y : Bool) : Bool = not(and(x,y))

// Lists

sealed trait List[+A]
case object Empty extends List[Nothing]
case class Cons[A](x : A, xs : List[A]) extends List[A]

def length[A](xs : List[A]) : Nat = xs match {
    case Empty => Zero
    case Cons(_, xs) => Succ(length(xs))
}

def append[A](xs : List[A], ys : List[A]) : List[A] = xs match {
    case Empty => ys
    case Cons(x, xs) => Cons(x, append(xs, ys))
}

def reverse[A](xs : List[A]) : List[A] = xs match {
    case Empty => Empty
    case Cons(x, xs) => append(xs, Cons(x, Empty))
}

def map[A,B](f : (A => B), xs : List[A]) : List[B] = xs match {
    case Empty => Empty
    case Cons(x, xs) => Cons(f(x), map(f,xs))
}

// TEST HELPER
def passed(points: Int) {
    require(points >=0)
    if (points == 1) print(s"\n*** Tests Passed (1 point) ***\n")
    else print(s"\n*** Tests Passed ($points points) ***\n")
}


```

Output:
```
defined trait Nat
defined object Zero
defined class Succ
one: Succ = Succ(Zero)
two: Succ = Succ(Succ(Zero))
three: Succ = Succ(Succ(Succ(Zero)))
four: Succ = Succ(Succ(Succ(Succ(Zero))))
five: Succ = Succ(Succ(Succ(Succ(Succ(Zero)))))
six: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Zero))))))
seven: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero)))))))
eight: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero))))))))
nine: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero)))))))))
ten: Succ = Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Succ(Zero))))))))))
plus: (x: Nat, y: Nat)Nat
nat_to_int: (x: Nat)Int
print_nat: (x: Nat)String
mult: (x: Nat, y: Nat)Nat
defined trait Bool
defined object True
defined object False
t: True.type
f: False.type
id: (x: Bool)Bool
not: (x: Bool)Bool
and: ...
```

# Homework 2

Due 9/12 at 11:59pm

## Submission Instructions
Upload only this .ipynb file to Canvas. Do not add anything to hw2stdlib since you can't submit it.

In this homework we will develop more useful functions for numbers as well as some of the functions on lists that we spoke about during lecture.

## Arithmetic Functions

### Problem 1 (5 Points)

Write the minus function, which subtracts a natural number from another natural number. If the result would be less than zero then leave the answer as zero.


<!-- grade_id: cell-19061b6ebc83ee13, points: 0, grade: false, solution: true -->
```scala
def minus(x : Nat, y : Nat) : Nat =
    (x, y) match{
        case (Zero, y) => Zero
        case (x, Zero) => x
        case (Succ(x), Succ(y)) => minus(x, y)
    }

```

Output:
```
minus: (x: Nat, y: Nat)Nat
```

<!-- grade_id: cell-80b015bfaabd7c53, points: 3, grade: true, solution: false -->
```scala
assert(minus(four, two) == two, 1)
assert(minus(four, Zero) == four, 2)
assert(minus(Zero, four) == Zero, 3)
passed(3)
```

Output:
```
*** Tests Passed (3 points) ***
```

<!-- grade_id: cell-3fb81b3d6c3bdcd6, points: 2, grade: true, solution: false -->
```scala
// HIDDEN TEST (2 pts)
```

Output:
```
input is incomplete
```

### Problem 2 (5 points)

Define the $\texttt{pow}$ function, which raises a natural number to some power. Notationally: 

$$
f\ x\ y = x^y
$$

Hint: Use the mult function from the stdlib for help

<!-- grade_id: cell-cf03509753bbf640, points: 0, grade: false, solution: true -->
```scala
def pow(x : Nat, y : Nat) : Nat = y match {
    case Zero => one
    case Succ(y) => mult(pow(x,y), x)
}
```

Output:
```
pow: (x: Nat, y: Nat)Nat
```

<!-- grade_id: cell-cc0a85f98f9550b1, points: 4, grade: true, solution: false -->
```scala
assert(pow(three, two) == nine, 1)
assert(pow(four, Zero) == one, 2)
assert(pow(two, two) == four, 3)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-e4f7652875d80790, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
```

Output:
```
input is incomplete
```

### Problem 3 (5 points)

Write the "equals to" function for numbers. This function should return $True$ if $x == y$ and $False$ otherwise.

<!-- grade_id: cell-5c44d445fb4cbec0, points: 0, grade: false, solution: true -->
```scala
def eq_nat(x : Nat, y : Nat) : Bool = (x, y) match {
    case (Zero, Zero) => True
    case (Succ(x), Succ(y)) => eq_nat(x, y)
    case (_, _) => False
}
```

Output:
```
eq_nat: (x: Nat, y: Nat)Bool
```

<!-- grade_id: cell-cc0a85f98f9550b2, points: 4, grade: true, solution: false -->
```scala
assert(eq_nat(four, four) == True, 1)
assert(eq_nat(four, Zero) == False, 2)
assert(eq_nat(eight, nine) == False, 3)
assert(eq_nat(nine, nine) == True, 4)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-e4f7652875d80791, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
```

Output:
```
input is incomplete
```

## Functions on Number Lists

### Problem 4 (5 pts)

Write the sum function. This function should take a list of natural numbers and return their sum. Here is the type:

$$
sum : \text{List}\ \mathbb{N} \rightarrow \mathbb{N}
$$

Hint: An empty list should have sum 0.

<!-- grade_id: cell-1026d25f142fc212, points: 0, grade: false, solution: true -->
```scala
def sum[A](xs: List[Nat]) : Nat = xs match {
    case Empty => Zero
    case Cons(x, xs) => plus(x, sum(xs))
}
```

Output:
```
sum: [A](xs: List[Nat])Nat
```

<!-- grade_id: cell-c8396ec50b153d54, points: 4, grade: true, solution: false -->
```scala
assert(sum(Cons(one, Empty)) == one, 1)
assert(sum(Cons(one, Cons(seven, Cons(two, Empty)))) == ten, 2)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
```

<!-- grade_id: cell-d118cae84cab07ee, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
```

Output:
```
input is incomplete
```

### Problem 5 (5 pts)

Write the product function. This function should take alist of natural numbers and return their product. Here is the type:

$$
prod : \text{List}\ \mathbb{N} \rightarrow \mathbb{N}
$$

Hint: An empty list should have product 1.

<!-- grade_id: cell-6a2caa631bc53781, points: 0, grade: false, solution: true -->
```scala
def prod[A](xs: List[Nat]) : Nat = xs match {
    case Empty => one
    case Cons(x, xs) => mult(x, prod(xs))
}
```

Output:
```
prod: [A](xs: List[Nat])Nat
```

<!-- grade_id: cell-cc0a85f98f9550b3, points: 4, grade: true, solution: false -->
```scala
assert(prod(Cons(one, Empty)) == one, 1)
val fourteen = Succ(Succ(Succ(Succ(ten))))
assert(prod(Cons(one, Cons(seven, Cons(two, Empty)))) == fourteen, 2)
passed(4)
```

Output:
```
*** Tests Passed (4 points) ***
null
```

<!-- grade_id: cell-e4f7652875d80792, points: 1, grade: true, solution: false -->
```scala
// HIDDEN TEST (1 pt)
```

Output:
```
input is incomplete
```

### Problem 6 (6 pts)

Write the function that tells us if a given natural number is in a list of natural numbers. (you will need to use the equals function defined in problem 3)

$$
in : \mathbb{N} \rightarrow \text{List}\ \mathbb{N} \rightarrow \mathbb{B}
$$

<!-- grade_id: cell-603f77bab9b8f27a, points: 0, grade: false, solution: true -->
```scala
def in(x : Nat, xs : List[Nat]) : Bool = xs match
{
    case Empty => False
    case Cons(n, ns) => eq_nat(x, n) match
    {
        case True => True
        case False => in(x, ns)
    }
}
```

Output:
```
in: (x: Nat, xs: List[Nat])Bool
```

<!-- grade_id: cell-b9d72e55abc06dd4, points: 6, grade: true, solution: false -->
```scala
assert(in(one, Cons(one, Empty)) == True, 1)
assert(in(seven, Cons(one, Cons(seven, Cons(two, Empty)))) == True, 2)
assert(in(eight, Cons(one, Cons(seven, Cons(two, Empty)))) == False, 3)
assert(in(seven, Empty) == False, 4)
passed(6)
```

Output:
```
*** Tests Passed (6 points) ***
```

```scala

```